<img src='https://hammondm.github.io/hltlogo1.png' style="float:right">

Linguistics 578<br>
Fall 2024<br>
Hammond

## Things to remember about any homework assignment:

1. For this assignment, you will edit this jupyter notebook and turn it in. Do not turn in pdf files or separate `.py` files.
1. Late work is not accepted.
1. Given the way I grade, you should try to answer *every* question, even if you don't like your answer or have to guess.
1. You may *not* use `python` modules that we have not already used in class.
1. You may certainly talk to your classmates about the assignment, but everybody must turn in *their own* work. It is not acceptable to turn in work that is essentially the same as the work of classmates.
1. All code must run. It doesn't have to be perfect, it may not do all that you want it to do, but it must run without error.
1. Code must run in reasonable time. Assume that if it takes more than *5 minutes* to run (on your machine), that's too long.
1. Please do not add, remove, or copy autograded cells.
1. Make sure to select `restart, run all cells` from the `kernel` menu when you're done and before you turn this in!

***

***my name***: [put your name here]

***people I talked to about the assignment***: [put your answer here]

***

## Homework #2

Here are the imports. Please do not import anything else.

```python
import numpy as np
from scipy.io import wavfile
from numpy.fft import rfft,irfft,rfftfreq
import matplotlib.pyplot as plt
import librosa
```

1. Write a function that will break a sound file into a specific number of windows.

<!-- grade_id: q1a, points: 0, grade: false, solution: true -->
```python
def makewindows(w,n):
    '''break a wave file into windows
    args:
        w: the sound file (as a numpy array)
        n: a specific number of windows
    returns:
        the windows as a numpy array/matrix.
        (each row is a window)
    '''
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: q1t1, points: 1, grade: true, solution: false -->
```python
res = makewindows(np.arange(50),8)
assert res.shape == (8,6)
```

<!-- grade_id: q1t2, points: 1, grade: true, solution: false -->
```python
assert len(res[7]) == 6
```

<!-- grade_id: q1t3, points: 1, grade: true, solution: false -->
```python
assert res[3,4] == 22
```

2. Imagine you have a sound file that is 2.2 seconds long and sampled at 16000 samples a second that you break into windows that are .03sec long. How many samples occur in each window? (Round to the nearest whole number.)

<!-- grade_id: q2a, points: 0, grade: false, solution: true -->
```python
#x = ?
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: q2t1, points: 1, grade: true, solution: false -->
```python
#this test is hidden!
```

3. Using the file `mhi.wav`, compute a cepstrum over the whole file and save it as `iceps`. Make sure to load the wave with `wavfile.read()`.

<!-- grade_id: q3a, points: 0, grade: false, solution: true -->
```python
#define these variables:
#fs,w = ?
#...
#iceps = ?
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: q3t1, points: 1, grade: true, solution: false -->
```python
assert type(iceps) == np.ndarray
```

<!-- grade_id: q3t2, points: 1, grade: true, solution: false -->
```python
assert np.isclose(iceps.shape[0],117626,atol=5)
```

Here we plot the lower cepstrum values:

```python
plt.plot(np.arange(10,400),np.abs(iceps[10:400]))
plt.show()
```

4. What is the frequency `f` and quefrency `q` of the fundamental frequency?

<!-- grade_id: q4a, points: 0, grade: false, solution: true -->
```python
#q = ?
#f = ?
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: q4t1, points: 1, grade: true, solution: false -->
```python
#this test is hidden!
```

<!-- grade_id: q4t2, points: 1, grade: true, solution: false -->
```python
#this test is hidden!
```

5. With the wavefile you loaded for question \#3, take the samples from 1000 to 1500 and compute a 4th-order LPC.

<!-- grade_id: q5a, points: 0, grade: false, solution: true -->
```python
#lpc = ?
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: q5t1, points: 1, grade: true, solution: false -->
```python
assert lpc.shape[0] == 5
```

<!-- grade_id: q5t2, points: 1, grade: true, solution: false -->
```python
assert lpc[0] == 1
```

<!-- grade_id: q5t3, points: 1, grade: true, solution: false -->
```python
assert np.isclose(lpc[1],-0.69,atol=.1)
```