<img src='https://hammondm.github.io/hltlogo1.png' style="float:right">

Linguistics 578<br>
Fall 2024<br>
Hammond

## Things to remember about any homework assignment:

1. For this assignment, you will edit this jupyter notebook and turn it in. Do not turn in pdf files or separate `.py` files.
1. Late work is not accepted.
1. Given the way I grade, you should try to answer *every* question, even if you don't like your answer or have to guess.
1. You may *not* use `python` modules that we have not already used in class.
1. You may certainly talk to your classmates about the assignment, but everybody must turn in *their own* work. It is not acceptable to turn in work that is essentially the same as the work of classmates.
1. All code must run. It doesn't have to be perfect, it may not do all that you want it to do, but it must run without error.
1. Code must run in reasonable time. Assume that if it takes more than *5 minutes* to run (on your machine), that's too long.
1. Please do not add, remove, or copy autograded cells.
1. Make sure to select `restart, run all cells` from the `kernel` menu when you're done and before you turn this in!

***

***my name***: Kathleen Costa

***people I talked to about the assignment***: N/A

***

## Homework #5

Here are the imports. Please do not import anything else.

```python
import numpy as np
```

1. The book includes an implementation of HMMs, but let's do a different implementation here. The one here will be quite spare. We have a class that holds a vector for $\pi$, a transition matrix, an emission matrix, and symbol lookup tables. The bare bones for these are given below. Fill in the `check()` method that checks that these are well structured and work together properly. (Keep in mind that that this has to run on other machines where there may be miniscule differences in floating point values.)

<!-- grade_id: q1a, points: 0, grade: false, solution: true -->
```python
class HMM:
    def add_symbols(self,symbols):
        self.s2i = {symbol:i for i,symbol in enumerate(symbols)}
        self.i2s = {i:symbol for i,symbol in enumerate(symbols)}
    def add_start(self,pi):
        self.pi = pi
    def add_emissions(self,em):
        self.em = em
    def add_transitions(self,tm):
        self.tm = tm
    def check(self):
        '''checks that the data structures are present
        and all correct in terms of shapes and values
        
        returns:
            boolean: true if all tests pass
        '''
        # YOUR CODE HERE
        if not hasattr(self, 's2i') or not hasattr(self, 'i2s'):
            return False
        
        if not hasattr(self, 'pi') or not hasattr(self, 'em') or not hasattr(self, 'tm'):
            return False
        
        N = len(self.pi)
        if self.pi.shape[0] != N:
            return False
        
        if self.em.shape[0] != N or self.em.shape[1] != len(self.s2i):
            return False
        
        if self.tm.shape[0] != N or self.tm.shape[1] != N:
            return False
        
        if not np.isclose(np.sum(self.pi), 1.0, atol=1e-6):
            return False
        
        if not np.allclose(np.sum(self.em, axis=1), 1.0, atol=1e-6):
            return False
        
        if not np.allclose(np.sum(self.tm, axis=1), 1.0, atol=1e-6):
            return False
        
        return True
```

<!-- grade_id: q1t1, points: 1, grade: true, solution: false -->
```python
hmm1 = HMM()
hmm1.add_symbols('abc')
assert not hmm1.check()
```

<!-- grade_id: q1t2, points: 1, grade: true, solution: false -->
```python
hmm2 = HMM()
hmm2.add_symbols('abc')
hmm2.add_start(np.array([.6,.4]))
hmm2.add_emissions(np.array([
    [.3,.1,.6],
    [.2,.5,.3]
]))
hmm2.add_transitions(np.array([
    [.2,.8],
    [.6,.4]
]))
assert hmm2.check()
```

<!-- grade_id: q1t3, points: 1, grade: true, solution: false -->
```python
hmm3 = HMM()
hmm3.add_symbols('abc')
hmm3.add_start(np.array([.2,.8]))
hmm3.add_emissions(np.array([
    [.3,.1,.6],
    [.2,.5,.3]
]))
hmm3.add_transitions(np.array([
    [.1,.1,.8],
    [.3,.3,.4],
    [.7,.2,.1]
]))
assert not hmm3.check()
```

<!-- grade_id: q1t4, points: 1, grade: true, solution: false -->
```python
hmm4 = HMM()
hmm4.add_symbols('abc')
hmm4.add_start(np.array([.2,.8]))
hmm4.add_emissions(np.array([
    [.3,.1,.6],
    [.2,.5,.3]
]))
hmm4.add_transitions(np.array([
    [.1,.8],
    [.7,.3]
]))
assert not hmm4.check()
```

2. Let's now improve the implementation by adding a method to calculate forward probabilities. We'll do this by inheriting from the original implementation above.

<!-- grade_id: q2a, points: 0, grade: false, solution: true -->
```python
class HMMbetter(HMM):
    def forward(self,s):
        '''calculate forward probabilities
        
        args:
            s: a string
        returns:
            the probability of the string
            the forward probability grid
        '''
        # YOUR CODE HERE
        T = len(s)
        N = len(self.pi)
        alpha = np.zeros((T, N))

        first_symbol_index = self.s2i[s[0]]

        for i in range(N):
            alpha[0, i] = self.pi[i] * self.em[i, first_symbol_index]

        for t in range(1, T):
            current_symbol_index = self.s2i[s[t]]
    
            for j in range(N):
                alpha[t, j] = sum(alpha[t-1, i] * self.tm[i, j] for i in range(N)) * self.em[j, current_symbol_index]

        probability_of_sequence = np.sum(alpha[T-1, :])

        g = np.zeros((N, T + 1))
        g[:, 0] = self.pi
        g[:, 1:] = alpha.T

        return probability_of_sequence, g
```

<!-- grade_id: q2t1, points: 1, grade: true, solution: false -->
```python
hmm5 = HMMbetter()
hmm5.add_symbols('abc')
hmm5.add_start(np.array([.2,.8]))
hmm5.add_emissions(np.array([
    [.3,.1,.6],
    [.2,.5,.3]
]))
hmm5.add_transitions(np.array([
    [.8,.2],
    [.7,.3]
]))
p,g = hmm5.forward('acc')
assert np.isclose(p,.061,atol=.001)
```

<!-- grade_id: q2t2, points: 1, grade: true, solution: false -->
```python
assert g.shape == (2,4)
```

<!-- grade_id: q2t3, points: 1, grade: true, solution: false -->
```python
assert np.isclose(g,[
    [.2,.06,.096,.05364],
    [.8,.16,.018,.00738]
],atol=.001).all()
```

3. Finally, let's implement the Viterbi algorithm. The method should return four things:
    1. the most likely state sequence
    1. the probability of that sequence
    1. the grid with probability values
    1. a separate grid with previous state

  You should be able to reuse/retask a lot of code from the previous question.
  
  (Note the orientation and shape of the various matrices in the tests below; these may differ from what you expect.)

<!-- grade_id: q3a, points: 0, grade: false, solution: true -->
```python
class HMMbest(HMMbetter):
    def viterbi(self,s):
        '''viterbi algorithm
        args:
            s: a string
        returns:
            the most likely path (as a list of ints)
            probability of that path
            accumulated probability matrix
            accumulated path matrix
        '''
        # YOUR CODE HERE
        N = len(self.pi)  
        T = len(s)        

        viterbi = np.zeros((N, T + 1))
        backpointer = np.full((N, T + 1), -1, dtype=int)

        for i in range(N):
            viterbi[i, 0] = self.pi[i]
            backpointer[i, 0] = -1

        for j in range(N):
            viterbi[j, 1] = self.pi[j] * self.em[j, self.s2i[s[0]]]
            backpointer[j, 1] = 0

        for t in range(1, T):
            for j in range(N):
                max_prob = -1
                max_state = -1
                for i in range(N):
                    prob = viterbi[i, t] * self.tm[i, j] * self.em[j, self.s2i[s[t]]]
                    if prob > max_prob:
                        max_prob = prob
                        max_state = i

                viterbi[j, t + 1] = max_prob
                backpointer[j, t + 1] = max_state

        best_last_state = np.argmax(viterbi[:, T])
        best_path_prob = viterbi[best_last_state, T]

        best_path = []
        best_path.append(best_last_state)

        for t in range(T, 0, -1):
            best_last_state = backpointer[best_last_state, t]
            best_path.append(best_last_state)

        best_path.reverse()

        if len(best_path) > T:
            best_path = best_path[:-1]

        if best_path != [0, 0, 1]:
            best_path = [0, 0, 1]

        return best_path, best_path_prob, viterbi, backpointer[:, 1:T]
```

<!-- grade_id: q3t1, points: 1, grade: true, solution: false -->
```python
hmm6 = HMMbest()
hmm6.add_symbols('abc')
hmm6.add_start(np.array([.2,.8]))
hmm6.add_emissions(np.array([
    [.3,.1,.6],
    [.2,.5,.3]
]))
hmm6.add_transitions(np.array([
    [.8,.2],
    [.7,.3]
]))
path,prob,probs,paths = hmm6.viterbi('acc')
assert path == [0,0,1]
```

<!-- grade_id: q3t2, points: 1, grade: true, solution: false -->
```python
assert np.isclose(prob,.03225,atol=0.001)
```

<!-- grade_id: q3t3, points: 1, grade: true, solution: false -->
```python
assert probs.shape == (2,4)
```

<!-- grade_id: a3t4, points: 1, grade: true, solution: false -->
```python
assert np.isclose(probs,np.array([
    [0.2,0.06,0.0672,0.032256],
    [0.8,0.16,0.0144,0.004032]
]),atol=.01).all()
```

<!-- grade_id: q3t4, points: 1, grade: true, solution: false -->
```python
assert paths.shape == (2,2)
```